---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Primero, se puede trabajar en R como si fuera una calculadora. Con esto se pueden realizar diversos cálculos.			

`@instructions`
Realizar la suma de 1 con 2.

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
###Suma
```

`@solution`
```{r}
1+2
```

`@sct`
```{r}

```
